### Week 6-1 NumPy Practice Problems

```python
import numpy as np
```

### Warm up
NumPy creation and indexing.

```python
# 1. Creates 3x3 array of zeros
z = np.zeros((3,3))
z
```

Output:
```
array([[0., 0., 0.],
       [0., 0., 0.],
       [0., 0., 0.]])
```

```python
# 2. Creates array [0, 2, 4, 6, 8]
a = np.array([0,2,4,6,8])
a
```

Output:
```
array([0, 2, 4, 6, 8])
```

```python
# 3. 5 evenly spaced numbers from 0 to 1
s = np.linspace(0,1,5)
s
```

Output:
```
array([0.  , 0.25, 0.5 , 0.75, 1.  ])
```

```python
# Array indexing, Gets elements > 5
arr = np.array([[1, 2, 3], 
                   [4, 5, 6], 
                   [7, 8, 9]])

arr_l5 = arr[arr<5]
arr_l5
```

Output:
```
array([1, 2, 3, 4])
```

### Problem 1
Convert an array of Celsius measurements to Fahrenheit.
Note: `fahrenheit = (Celsius * 9/5) + 32`.

```python
def convert_to_fahrenheit_list(temperatures):
    cels = []
    for i in range(0,len(temperatures)):
        cels.append((temperatures[i] * 9/5) + 32)

    return cels
```

```python
print("Testing with a regular list:")
temps_list = [0, 10, 20, 30]
print(convert_to_fahrenheit_list(temps_list))
```

Output:
```
Testing with a regular list:
[32.0, 50.0, 68.0, 86.0]
```

```python
def convert_to_fahrenheit_numpy_array(temperatures):
    temp_array = np.array(temperatures)

    return (temp_array * 9/5) + 32
```

```python
print("\nTesting with a NumPy array:")
temps_array = np.array([0, 10, 20, 30])
print(convert_to_fahrenheit_numpy_array(temps_array))
```

Output:
```
Testing with a NumPy array:
[32. 50. 68. 86.]
```

### Problem 2
Use numpy to create a function that adds a border filled with zeros around an existing array.

*Hint*: Use numpy's `pad` function. See https://numpy.org/doc/stable/reference/generated/numpy.pad.html

Note: we did not cover this function in class, see if you are comfortable of reading numpy documents and learn on your own!

<!-- grade_id: cell-2466342b05a17106, points: 0, grade: false, solution: true -->
```python
def add_border(x):
    x_arr = np.array(x)
    return np.pad(x_arr, (1,1), "constant", constant_values = 0)
```

<!-- grade_id: cell-fb9bae34e3908aed, points: 1, grade: true, solution: false -->
```python
x = np.ones((1, 1))
claimed = add_border(x)
correct = np.zeros((3, 3))
correct[1, 1] = 1.0
assert claimed.shape == correct.shape
assert np.all(np.abs(claimed - correct) < 1e-8)
```

```python
add_border(np.array(((1,2),(1,2))))
```

Output:
```
array([[0, 0, 0, 0],
       [0, 1, 2, 0],
       [0, 1, 2, 0],
       [0, 0, 0, 0]])
```

### Problem 3
Create an $n \times n$ numpy array filled with a checkerboard pattern of zeros and ones. The first entry should be the integer `0`. Assume $n \ge 2$.

```python
np.reshape(np.tile([0,1], 2), shape = (2,2))
```

Output:
```
array([[0, 1],
       [0, 1]])
```

```python
n = 5
first = np.tile([0,1], int(n/2))
second = np.tile([1,0], int(n/2))
```

```python
first_o = np.append(first,0)
second_o = np.append(second, 1)
almost = np.tile(np.array([first_o, second_o]), (int(n/2), 1))
np.append(almost, [first_o], axis = 0)
```

Output:
```
array([[0, 1, 0, 1, 0],
       [1, 0, 1, 0, 1],
       [0, 1, 0, 1, 0],
       [1, 0, 1, 0, 1],
       [0, 1, 0, 1, 0]])
```

```python
almost
```

Output:
```
array([[0, 1, 0, 1, 0],
       [1, 0, 1, 0, 1],
       [0, 1, 0, 1, 0],
       [1, 0, 1, 0, 1]])
```

```python
first_o
```

Output:
```
array([0, 1, 0, 1, 0])
```

<!-- grade_id: cell-9982d08b4ea7057d, points: 0, grade: false, solution: true -->
```python
def make_checkerboard(n):
    first = np.tile([0,1], int(n/2))
    second = np.tile([1,0], int(n/2))
    if n % 2 == 0:
        return np.tile(np.array([first, second]), (int(n/2), 1))
    if n % 2 != 0:
        first_o = np.append(first,0)
        second_o = np.append(second, 1)
        almost = np.tile(np.array([first_o, second_o]), (int(n/2), 1))
        return np.append(almost, [first_o], axis = 0)

```

```python
make_checkerboard(n = 5)
```

Output:
```
array([[0, 1, 0, 1, 0],
       [1, 0, 1, 0, 1],
       [0, 1, 0, 1, 0],
       [1, 0, 1, 0, 1],
       [0, 1, 0, 1, 0]])
```

<!-- grade_id: cell-e6f9eb1adf863ce2, points: 1, grade: true, solution: false -->
```python
correct = \
    [[0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0],
     [0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0],
     [0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0],
     [0, 1, 0, 1, 0, 1, 0, 1],
     [1, 0, 1, 0, 1, 0, 1, 0]]
assert np.all(correct == make_checkerboard(8))
```

```python
n = 6
z = np.zeros((n, n))
z[1::2,::2] = 1
z[::2, 1::2] = 1
```

```python
z
```

Output:
```
array([[0., 1., 0., 1., 0., 1.],
       [1., 0., 1., 0., 1., 0.],
       [0., 1., 0., 1., 0., 1.],
       [1., 0., 1., 0., 1., 0.],
       [0., 1., 0., 1., 0., 1.],
       [1., 0., 1., 0., 1., 0.]])
```

```python

```